# BT01: Ẩn tin mật trên văn bản bằng phương pháp chèn khoảng trắng <font color="grey">(chèn khoảng trắng giữa các từ để làm cho văn bản được căn lề 2 bên)


Họ tên: Hồ Trung Hiếu

MSSV: 19424003

---
## Cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT01`
        - File `BT01-AnTinMatTrenVanBan_KhoangTrang.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác cấu trúc này.</font>

## 2. Import

```python
from bitarray import bitarray
import math
```

## 3. Hàm nhúng (6đ)

<!-- grade_id: cell-2f21b430087c5c35, points: 0, grade: false, solution: true -->
```python
def embed(msg_file, cover_text_file, text_width, stego_text_file):
    '''
    Nhúng tin mật vào văn bản bằng phương pháp chèn khoảng trắng 
    để văn bản được căn lề 2 bên.
    
    Các tham số:
        msg_file (str): Tên file chứa secret message.
        cover_text_file (str): Tên file chứa cover text.
        text_width (int): Chiều dài của dòng sau khi được căn lề.
        stego_text_file (str): Tên file chứa stego text.
    Giá trị trả về:
        bool: True nếu nhúng thành không, False nếu không đủ chỗ để nhúng. 
    '''
    # Read cover text file
    with open(cover_text_file, 'r') as f:
        cover_text_lines = f.read().splitlines()
    
    # Read msg file
    with open(msg_file, 'r') as f:
        msg = f.read()
    
    # Convert msg to msg bits
    msg_bits = bitarray()
    msg_bits.fromstring(msg)
    
    # Embed msg bits into cover text lines
    stego_text = ''
    b = 0 #chi so tung bit cua meg_bits
    i=0 #chi so dong cua covertext
    countEnd = 0
    while i< len(cover_text_lines): # duyet tung dong
        listchar = list(cover_text_lines[i])#list cac tu o line i thanh cac ky tu
        idx_space = [temp for temp in range(len(cover_text_lines[i])) if cover_text_lines[i][temp] == " " ]
        if len(cover_text_lines[i]) <= text_width: #dong co the can le
            n = len(cover_text_lines[i].split()) -1 #so khoang trong cua dong i(in line)
            m = text_width - len(cover_text_lines[i]) #so khoang trang de dong i can le(cuoi dong)
            print(i + 1, m, n)
            if 0 < m and m < n:
                x = 0
                count = 0
                while count < min(m, n-m):#duyệt min cap khoang trang->so khong trang chen duoc trong dong dau tien tu 0->min
                    if b < len(msg_bits):
                        if not msg_bits[b]:
                            listchar[int(idx_space[x])] = '  '#chen bit 0
                            x+=2
                        else:
                            x += 1
                            listchar[int(idx_space[x])] = '  '#chen bit 1
                            x += 1
                        b += 1
                    else:
                        if countEnd == 0:
                            x += 1
                            listchar[int(idx_space[x])] = '  '#nhung 1
                            x += 1
                            countEnd +=1
                        else:
                            listchar[int(idx_space[x])] = '  '#nhung 0
                            x += 2
                    count+=1
                if min(m, n-m) == n-m: #cac khoang trang con lai
                    runF = n-m
                    for temp in range(runF, n):
                        #print(temp)
                        listchar[int(idx_space[temp])] = '  '
                        #print('orther', i, b, temp, idx_space[temp], ''.join(listchar))
                listchar = ''.join(listchar)
                stego_text += listchar
            if m > n:
                listchar = ''.join(listchar)
                stego_text += listchar
            if m == n:
                for temp in range(0, m):
                    listchar[int(idx_space[temp])] = '  '
                listchar = ''.join(listchar)
                stego_text += listchar
                
            if i + 1 != len(cover_text_lines):
                    stego_text += '\n'
        i+=1
    
    if b < len(msg_bits):
        print('Not enough spaces to embed!')
        return False
    
    # Write stego text to file
    with open(stego_text_file, 'w') as f:
        f.write(stego_text)
    return True

embed('msg.txt', 'cover.txt', 70, 'stego.txt')
```

Output:
```
1 57 1
2 70 -1
3 4 10
4 10 10
5 70 -1
6 0 10
7 9 9
8 4 7
9 13 6
10 70 -1
11 1 12
12 2 8
13 0 11
14 2 10
15 3 12
16 5 10
17 2 9
18 7 6
19 29 4
20 70 -1
21 4 9
22 4 10
23 1 9
24 8 9
25 12 7
26 70 -1
27 2 10
28 4 8
29 0 12
30 55 2
31 70 -1
32 1 7
33 0 7
34 1 10
35 6 10
36 1 8
37 0 12
38 3 11
39 2 11
40 30 7
41 70 -1
42 1 12
43 1 9
44 1 12
45 3 9
46 1 12
47 4 11
48 2 8
49 0 13
50 0 10
51 2 9
52 70 -1
53 2 8
54 0 10
55 5 6
56 7 7
57 60 0
58 70 -1
59 67 0
True
```

<!-- grade_id: cell-057a58bf038e1e60, points: 5, grade: true, solution: false -->
```python
# TEST 1            
result = embed('msg.txt', 'cover.txt', 70, 'stego.txt')
assert result == True
with open('stego.txt', 'r') as f:
    stego_text = f.read()
with open('correct_stego.txt', 'r') as f:
    correct_stego_text = f.read()
assert stego_text == correct_stego_text
```

Output:
```
4 10
4 7
orther 7 7 3 32 steganography  combines the  Greek  words steganos,  meaning "covered,
orther 7 7 4 38 steganography  combines the  Greek  words  steganos,  meaning "covered,
orther 7 7 5 48 steganography  combines the  Greek  words  steganos,  meaning "covered,
orther 7 7 6 56 steganography  combines the  Greek  words  steganos,  meaning  "covered,
1 12
2 8
2 10
3 12
5 10
orther 15 20 5 39 invisible ink  between  the visible  lines  of a  private  letter. Some
orther 15 20 6 42 invisible ink  between  the visible  lines  of  a  private  letter. Some
orther 15 20 7 44 invisible ink  between  the visible  lines  of  a  private  letter. Some
orther 15 20 8 52 invisible ink  between  the visible  lines  of  a  private  letter. Some
orther 15 20 9 60 invisible ink  between  the visible  lines  of  a  private  letter.  Some
2 9
4 9
4 10
1 9
8 9
orther 23 32 1 16 unbreakable they  are, arouse interest and may in themselves be
orther 23 32 2 21 unbreakable they  are,  arouse interest
```

<!-- grade_id: cell-f625774497df6542, points: 1, grade: true, solution: false -->
```python
# TEST 2
result = embed('msg2.txt', 'cover.txt', 70, 'stego.txt')
assert result == False
```

Output:
```
Not enough spaces to embed!
```

## 4. Hàm rút trích (4đ)

<!-- grade_id: cell-f3f9b34770f282aa, points: 0, grade: false, solution: true -->
```python
def extract(stego_text_file, extr_msg_file):
    '''
    Hàm rút trích tin mật đã được nhúng bằng phương pháp chèn khoảng trắng.
    
    Các tham số:
        stego_text_file (str): Tên file chứa stego text.
        extr_msg_file (str): Tên file chứa secret message được rút trích.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-caf1afe26d95b670, points: 4, grade: true, solution: false -->
```python
# TEST
extract('correct_stego.txt', 'extr_msg.txt')
with open('extr_msg.txt', 'r') as f:
    extr_msg = f.read()
with open('correct_extr_msg.txt', 'r') as f:
    correct_extr_msg = f.read()
assert extr_msg == correct_extr_msg
```